// variables5.rs
// Make me compile! Execute the command `rustlings hint variables5` if you want a hint :)

// I AM NOT DONE

fn main() {
    let number = "T-H-R-E-E"; // don't change this line
    println!("Spell a Number: {}", number);
    let number_int = 3; // Use a new variable for the integer value
    println!("Number plus two is: {}", number_int + 2);
}